# Boolean expressions
The following uses *Boolean* operations. What do you think, the following snippets are doing? Check by executing. 

```python
mylist = []
for k in range(10):
    if not (k % 2 == 0 or k % 3 == 0):
        mylist.append(k)
print(mylist)
# loop through each value from 0 to 9 inclusive
# if it cannot be evenly divided by 2 or 3 then add to list
```

Output:
```
[1, 5, 7]
```

```python
if 5:
    print("5 is true")
else:
    print("5 is false")
# "if 5:" is essentially "5 is True"
# if 0 goes to false
```

Output:
```
5 is true
```

What changes, if you replace `5` by `0`?

```python
mylist = []
for k in range(10):
    if (k % 2) and (k % 3):
        mylist.append(k)
print(mylist)
# if not 0 is the same as True
# if 0 is the same as False
# if False both times then add to list
# only returning a list of items that are not divisible by 2 or 3
```

Output:
```
[1, 5, 7]
```

Explain, how this works!

# Divisible by three

Write a function `divisible_by_3(n)` which has the following properties: If you call it with a number which can be divided by three, it should print out `The number xxx can be divided by three`, else it prints `The number xxx can not be divided by three`, where `xxx` is replaced by the argument with which the function was called. 

<!-- grade_id: cell-bfebc8f10afd614d, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
def divisible_by_3(n):
    if not (n % 3):
        return 'The number ' + str(n) + ' can be divided by three'
    else:
        return 'The number ' + str(n) + ' can not be divided by three'
```

Output:
```
'The number 9999 can be divided by three'
```

Check that your function works, by modifiying the following cell:

```python
divisible_by_3(132)
```

Output:
```
'The number 132 can be divided by three'
```

# Prime numbers

A number is a *prime* number, if it can only be divided by itself and the number 1. Write a function `is_prime(n)` which `return`s `True` if a number is prime and `False` if this is not the case. 

<!-- grade_id: cell-44a5b0d61b3535b5, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
def is_prime(n):
    out = []
    for i in range(1, n+1):
        if not (n % i):
            out += [i]
    if len(out) > 2:
        return False
    return True
```

Check that your program recognizes the numbers 13, 127, and $M_{19} = 2^{19}-1$ as prime numbers and $M_{67} = 2^{67}-1$ as composite. (Numbers of the type $M_p = 2^p-1$ are called *Mersenne* numbers.  They have many interesting properties and connections to perfect numbers. For more reading, see http://primes.utm.edu/mersenne/ )


<!-- grade_id: cell-55a0d49ceb34f02e, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
def Mersenne(n):
    return 2**n - 1
```

Output:
```
True
```

# Find all factors of a number

Write a function `print_factors(n)` which prints out a list of all factors of a number.

<!-- grade_id: cell-074d7dbf54d2028d, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
def print_factors(n):
    out = []
    for i in range(1,n+1):
        if not (n%i):
            out += [i]
    return out
```

Test with the following lines. (Does this fulfill your expectations?)

```python
print_factors(28)
print_factors(496)
print_factors(8128)
print_factors(13)
print_factors(127)
print_factors((2**19-1))
```

Output:
```
[1, 524287]
```

Modify this function to also calculate the sum of its factors (exclude the number itself from the list of factors)

<!-- grade_id: cell-4df088cd9b8715a8, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
def print_factors(n):
    out = []
    out_sum = 0
    for i in range(1,n+1):
        if not (n%i):
            out += [i]
    for j in out[:-1]:
        out_sum += j
    return out, out_sum
```

Output:
```
([1, 13], 14)
```

If the sum of factors equals the number itself, then this number is called a *perfect* number. Modify your function, so that it checks for a number being perfect. 


<!-- grade_id: cell-9de32979eeb8f29e, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
def print_factors(n):
    out = []
    out_sum = 0
    k = False
    for i in range(1,n+1):
        if not (n%i):
            out += [i]
    for j in out[:-1]:
        out_sum += j
    if n == out_sum:
        k = True
    return out, out_sum, k
```

Check that your program recognizes the number 28, 496 and 8128 as perfect numbers. (In case you happen to find an *odd* perfect number, you would have solved one of the oldest problems in number theory. )

<!-- grade_id: cell-2e365d619f605b58, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
print_factors(28)
print_factors(496)
print_factors(8128)
```

Output:
```
([1, 2, 4, 8, 16, 32, 64, 127, 254, 508, 1016, 2032, 4064, 8128], 8128, True)
```

Check that $2^{18}(2^{19}-1)$ is a perfect number! (This is slightly more challenging, and can easily give a memory error or require a long time to complete. Rethink, how you can calculate the factors of a number. Notice that if you find one factor you get a second one without searching.)

<!-- grade_id: cell-679aca27513c9a91, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
n = (2**18)*(2**19 - 1)
#print_factors(n)
print_factors((2**19-1))
print_factors(2**18)
# find them separately and see if the second part of the output is equal to n
```

Output:
```
137438691328
```

# Wallis product

The *Wallis* sequence is defined through
\begin{equation}
  W(n) = \prod_{k=1}^n \frac{2k}{2k-1} \;\frac{2k}{2k+1}
\end{equation}
a) Write a Python function `W(n)` which returns the n'th element of the Wallis sequence. Calculate W(1) and W(10). 

<!-- grade_id: cell-9d9450f8e010d06f, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
def W(n):
    out = 1
    for k in range(1, n+1):
        first = (2*k) / (2*k - 1)
        second = (2*k) / (2*k + 1)
        out *= (first * second)
    return out
W(1)
```

Output:
```
1.3333333333333333
```

<!-- grade_id: cell-0eafa1cbcf8fc31a, points: 0, grade: true, solution: false -->
```python
assert(abs(W(1) - 4/3)< 1.0e-16)
assert(abs(W(10) - 1.53385)< 1.0e-4)
```

b) It is known that 
\begin{equation}
\lim_{n\to \infty}  2 W(n)= \pi \approx 3.1415926535897932
\end{equation}
Roughly (up to a factor of 10) determine $n$, which produces the first three digits 3.14 of $\pi$ correctly. How high do you need to go to get the first six digits?

<!-- grade_id: cell-4cf3f8255b659f85, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
for i in range(30000, 500000, 10000):
    print(i, 2 * W(i))
# you need something around 30k to get 3.14159
```

Output:
```
30000 3.1415664741963476
40000 3.1415730189423545
50000 3.141576945822721
60000 3.1415795637565753
70000 3.1415814337160595
80000 3.141582836189326
90000 3.141583927003995
100000 3.1415847996570205
110000 3.141585513646858
120000 3.1415861086389474
130000 3.1415866120942066
140000 3.1415870436275823
150000 3.1415874176234904
160000 3.1415877448701357
170000 3.141588033617258
180000 3.1415882902814944
190000 3.1415885199285674
200000 3.1415887266109324
210000 3.1415889136093
220000 3.1415890836079097
230000 3.1415892388241113
240000 3.1415893811056357
250000 3.1415895120047135
260000 3.1415896328345707
270000 3.141589744714059
280000 3.1415898486022313
290000 3.1415899453256544
300000 3.141590035600933
310000 3.141590120052062
320000 3.14159019922495
330000 3.1415902735994807
340000 3.1415903435990145
350000 3.1415904095986042
360000 3.141590471931518
370000 3.1415905308951446
380000 3.1415905867554317
390000 3.1415906397511484
400000 3.1415906900969865
410000 3.141590737986932
420000 3
```

```python

```